# Homework 5 (DL Friday, October 27 at 12:00 PM)
#### ELEC-E8740 - Basics of sensor fusion - Autumn 2023


```python
import numpy as np
import matplotlib.pyplot as plt
```

## Question:

## Implement Gauss–Newton with line search for minimizing the cost function $J(x) = (1.1 − \sin(x))^2$. Use grid search with grid $\gamma \in [0,\frac{1}{10}, \frac{2}{10}, \ldots , 1]$. Hint: Beware of the singularity of the derivative at the minimum.

## Recall:
#### Check "Gauss–Newton Algorithm with Line Search" in page $15$ of houndout $5$. Here, we want to use the "Exact Line Search on Grid" method as in page $11$ of houndout $5$.
## Procedure:
####  First, you need to define measurement function $g(x)$, its Jacobian $G(x)$, and the cost function $J(x)$.
#### Consider the function $g(x)$ in this example as $g(x) = \sin(x)$ and the measurement as $y = 1.1$.

```python
def g(x):
    return np.sin(x)

y = 1.1
```

### Part a (1 point): In the section below, implement the code for the Jacobian, $G(x)$, and the cost function, $J(x)$. 

<!-- grade_id: cell-bddb9819c9ddc784, points: 0, grade: false, solution: true -->
```python
def G(x):
    """ 
    Implement the Jacobian of g(x).
    """
    # g'(x) = cos(x)
    return np.cos(x)
    # YOUR CODE HERE
    # raise NotImplementedError()
    
def J(x):
    """
    Implement the cost function.

    """
    return (1.1 - g(x))**2
    # YOUR CODE HERE
    # raise NotImplementedError()
```

<!-- grade_id: cell-feff75040d2db5e2, points: 1, grade: true, solution: false -->
```python
"""Check the result for several inputs"""
assert np.allclose(G(0.0), 1.0)
assert np.allclose(J(0.0), 1.21)
```

#### Before implementing Gauss-Newton with line search method, let's visually observe the minimums of the cost function $J(x)$, and the cost of an initial guess e.g. $x_0 = 3.3$. 


```python
x0 = 3.3 # initial guess
xs = np.linspace(-10, 10, 1000)
cost_results = [J(x) for x in xs]

plt.figure(figsize=(10,6))
plt.plot(xs, cost_results, linestyle='--', label='cost function')
plt.plot(x0, J(x0), 'r*', label='cost of initial x')
plt.legend()
plt.grid()
plt.xlabel('x')
plt.ylabel('cost')
plt.title('Cost function')
plt.show()
```

### Part b (1 point): Here we aim to implemet the "Exact Line Search on Grid" algorithm according to the Page $11$ of houndout $5$ (Alg.3).


<!-- grade_id: cell-1e5f436acb64286e, points: 0, grade: false, solution: true -->
```python
def Exact_LS_on_Grid(xi, del_x, Ng):
    '''
    Input:
        xi: result of the previous iteration
        del_x: the update direction
        Ng: the grid size
    Output:
        opt_gamma: optimal step size
    '''
    opt_gamma = 0 # Line 1 of Alg.3 (opt_gamma is gamma* in Alg.3)
    opt_cost = J(xi) # Line 1 of Alg.3 (opt_cost is J* in Alg.3)
    for j in range(0, Ng + 1): # Line 2 of Alg.3 (note that in the question gamma is starting from zero)
        # Implement exact line search method (line 3 to 8 of Alg.3).
        # You should call the cost function J(x) that you previously coded.
        # YOUR CODE HERE
        raise NotImplementedError()
    return opt_gamma
```

<!-- grade_id: cell-5823cd69763d6d76, points: 1, grade: true, solution: false -->
```python
"""Check the result for several inputs"""
assert np.allclose(Exact_LS_on_Grid(3.3 , -2, 10), 0.9)
assert np.allclose(Exact_LS_on_Grid(0.0 , 2, 10), 0.8)

```

### Part c (2 points): In the section below, implement Gauss Newton with Line Search method (page 15).


<!-- grade_id: cell-35f4613c4859ca67, points: 0, grade: false, solution: true -->
```python
def Gauss_Newton_LS(x_0, y, Ng, number_of_iterations):
    x = np.zeros((number_of_iterations + 1,)) # do not change this line.
    x[0] = x_0
    for i in range(number_of_iterations):
        
        # Find the update direction (Line 3 Alg.5).
        # for that you need to call functions g and G
        
        # YOUR CODE HERE
        raise NotImplementedError()
        
        # Find the optimal gamma by calling  Exact_LS_on_Grid function.
        # YOUR CODE HERE
        raise NotImplementedError()
        
        # Update x (Line 5 Alg.5). you should calculate x[i+1] from x[i]
        # YOUR CODE HERE
        raise NotImplementedError()
    return x
        
```

<!-- grade_id: cell-72db145139b0c1af, points: 1, grade: true, solution: false -->
```python

```

#### In the section below, we check the final result of the "Gauss_Newton_LS" function by considering several initial points, specifically $[-7.0, 0.0, 3.3]$, running for $5$ iterations, and employing a grid of size $10$.



<!-- grade_id: cell-3c9e66635b237f70, points: 1, grade: true, solution: false -->
```python
"""Check the result for several inputs"""
assert np.allclose(Gauss_Newton_LS(-7.0, 1.1, 10, 5)[-1], -4.6695, rtol=1e-03, atol=1e-04)
assert np.allclose(Gauss_Newton_LS(0.0, 1.1, 10, 5)[-1], 1.5603, rtol=1e-03, atol=1e-04)
assert np.allclose(Gauss_Newton_LS(3.3, 1.1, 10, 5)[-1], 1.5672, rtol=1e-03, atol=1e-04)

```

Feel free to change the initial point in the next section and observe the results through the values of cost function.

```python
initial_point = -3.3
xx = Gauss_Newton_LS(initial_point, 1.1, 10, 5)

plt.figure(figsize=(10,6))
plt.plot(xs, cost_results, linestyle='--', label='cost function')
plt.plot(initial_point, J(initial_point), 'g*', label='cost of initial x')
plt.plot(xx, J(xx), 'r--', label='LS method')
plt.legend()
plt.grid()
plt.xlabel('x')
plt.ylabel('cost')
plt.title('Cost function')
plt.show()
```